# Miniproyecto 1: Detección de globulos blancos

## Entrega 1: 
En todo proyecto en el campo de visión por computador una de las partes mas importantes es conocer los datos con los cuales se va a trabajar a lo largo de este. En terminos generales las bases de datos contienen todos los datos necesarios para poder solucionar la tarea segun dicho marco experimental. En esta primera entrega haremos la exploración de la base de datos, desde la apertura y lectura de las anotaciones como la apertura de imagenes. Para esto realizaremos 3 procesos, importación de archivos, revision de estadisticas y visualización de imagenes.

### Parte 1: Importación de archivos
La primera parte para poder explorar los archivos es tener la capacidad de importarlos, por esto haremos la importación de los diferentes archivos y datos que usarán en este miniproyecto. Esto es necesario para acceder a ellos, procesarlos y manejarlos de forma correcta para así poder realizar el procedimiento adecuadamente.

#### 1.1 Importación de imágenes
Aprender a hacer de forma efectiva la importación de imagenes es una parte fundamental, en especial cuando manejamos grandes bases de datos. Por esto, necesitamos aprender a realizarlo de forma automatica por lo que debemos estar en capacidad de obtener todos los nombres de las imagenes y así poder abrirlas de forma iterativa. Una librería que permite hacer estos procedimientos es glob(https://docs.python.org/es/3/library/glob.html) la cual dado cierto formato permite abrir todos los archivos que cumplan con esas resticciones. 
> **Consejo**: Para una explicación mas detallada de este procedimiento y como usar glob pueden referirse al laboratorio 0. 

Ahora bien, para abrir las imagenes inicialmente vamos a realizar el procedimiento de obtener las 3 carpetas posibles *train*, *valid* y *test*. Para esto deben completar el contenido de la siguiente celda, manteniendo el nombre de las variables y teniendo en cuenta las consideraciones a continuación: 

> - Tengan cuidado de obtener listas vacías, glob no lanza errores al encontrarse con que ningun archivo o carpeta cumple el formato, solo retorna una lista vacía y es uno de los errores mas comunes al usar la librería.  
> - Tengan cuidado al usar glob de no retornar una misma lista en dos ocasiones.
> - Recuerden mantener los nombres de las variables de las listas es importante siempre mantener la misma notación en especial cuando todas son del mismo tipo y tienen una forma similar.

```python
# Importación de las librerías necesarias.
# Si necesita usar más librerías agreguelas en este espacio para mantener su código organizado.
import glob
import os
import numpy as np

import json

import matplotlib.pyplot as plt
from skimage.io import imread
import cv2
```

<!-- grade_id: cell-60b9b6df553ba40b, points: 0, grade: false, solution: true -->
```python
list_train = glob.glob('data_mp1/BCCD/train/*') # variable que tendrá la lista de las imagenes de la carpeta de train
list_val =  glob.glob('data_mp1/BCCD/valid/*')  # variable que tendrá la lista de las imagenes de la carpeta de val
list_test =  glob.glob('data_mp1/BCCD/test/*')  # variable que tendrá la lista de las imagenes de la carpeta de test
# YOUR CODE HERE
```

<!-- grade_id: cell-c8c613173bd1926f, points: 0.2, grade: true, solution: false -->
```python
list_train.sort()
list_val.sort()
list_test.sort()

# pruebas basicas de funcionamiento correcto de la función
assert(len(list_train)!=0), 'Su lista esta vacia'
assert(len(list_val)!=0), 'Su lista esta vacia' 
assert(len(list_test)!=0), 'Su lista esta vacia' 
```

<!-- grade_id: cell-e4be25cb87e9ba4b, points: 0.2, grade: true, solution: false -->
```python
assert(np.isclose(len(list_train),224,0.1)), 'La cantidad de datos en su lista de entrenamiento es erroneo'
assert(np.isclose(len(list_val),65,0.1)), 'La cantidad de datos en su lista de validacion es erroneo'
assert(np.isclose(len(list_test),31,0.1)), 'La cantidad de datos en su lista de test es erroneo'
```

<!-- grade_id: cell-076a89fe57df81bf, points: 0.2, grade: true, solution: false -->
```python
assert (os.path.join('data_mp1','BCCD','train','BloodImage_00001_jpg.rf.ee79407f4ce4abd4fc56f42c176d324b.jpg') in list_train) \
or (os.path.join('.', 'data_mp1','BCCD','train','BloodImage_00001_jpg.rf.ee79407f4ce4abd4fc56f42c176d324b.jpg') in list_train), 'Su lista de entrenamiento no tiene los datos correctos'
assert (os.path.join('data_mp1','BCCD','valid','BloodImage_00000_jpg.rf.cf90db4f6c7327631ae0d822bab81d5f.jpg') in list_val) \
or (os.path.join('.', 'data_mp1','BCCD','valid','BloodImage_00000_jpg.rf.cf90db4f6c7327631ae0d822bab81d5f.jpg') in list_val), 'Su lista de validacion no tiene los datos correctos'
assert (os.path.join('data_mp1','BCCD','test','BloodImage_00038_jpg.rf.03231bab33fddc4fabbb488a9d458343.jpg') in list_test) \
or (os.path.join('.', 'data_mp1','BCCD','valid','BloodImage_00000_jpg.rf.cf90db4f6c7327631ae0d822bab81d5f.jpg') in list_val), 'Su lista de test no tiene los datos correctos'
```

Ahora que ya tenemos las rutas de todas las imagenes almacenadas en los 3 arreglos es importante poder abrir las imagenes como tal y tener la capacidad de graficarlas. Para este proceso podemos usar unicamente la librería matplotlib tanto para abrir la imagen y graficarla. En este caso vamos a realizar un subplot pequeño con 2 imagenes diferentes de la base de datos en la siguiente celda:


**Consejo**: Para una explicación mas detallada de este procedimiento y como abrir y graficar imagenes con matplotlib pueden referirse al laboratorio 0. Puede usar skimage.io.imread() para leer las imagenes.

> - Tengan cuidado con usar plt.show() es un error muy comun al trabajar en cuadernos de Jupyter ya que sin esto la figura no se mostrará al estar en un archivo de python y en este caso mostrará información extra.  
> - Recuerden que una imagen no debe tener ejes ya que estos no aportan información relevante, por esto busquen como poder apagar los ejes.
> - Cuando se realice una imagen con un unico canal deben especificar el mapa de color que se va a usar, el que es correcto para las imagenes, a no ser que contenga algun tipo de información especial es 'gray'
> - La función imread() recibe como input un path y retorna una matriz.

<!-- grade_id: cell-fd4c930512a2ede0, points: 0, grade: false, solution: true -->
```python
fig, ax = plt.subplots(1,2) 
# Usen esta notación para hacer los subplots 
# (recuerden que para acceder a cada elemento se puede hacer como si fuera un array: ax[idx])

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f64c3f291834abc3, points: 0.1, grade: true, solution: false -->
```python
assert(not ax[1].axison), 'No apago los ejes de la image 1'
assert(not ax[0].axison), 'No apago los ejes de la image 0'
```

#### 1.2 Importación de anotaciones
Como se habrán dado cuenta las anotaciones están en un formato que probablemente no hayan usado nunca, un archivo *.json*. Estos archivos contienen una gran cantidad de información, tanto de la base de datos como de las imagenes y las anotaciones en si. El primer paso para poder explorarlas es poder hacer la importación de estos archivos y abrir el diccionario que se encuentra en su interior. Para ello use la siguiente celda manteniendo los nombres de las variables (pueden **apoyarse** en el codigo y la explicacion en esta pagina: https://www.geeksforgeeks.org/json-load-in-python/).

> **recomendaciones:** Revisen en primer lugar que tipo de objeto es el archivo que acaban de abrir. Revisen cuales son sus elementos y que contiene cada uno de ellos. Observen que compone cada una de las partes de interés que tiene su archivo.

<!-- grade_id: cell-fdb89cfe009fb401, points: 0, grade: false, solution: true -->
```python
dict_train = '' # variable que tendrá el diccionario de las anotaciones de la carpeta de train
dict_val = '' # variable que tendrá el diccionario de las anotaciones de la carpeta de val
dict_test = '' # variable que tendrá el diccionario de las anotaciones de la carpeta de test
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# Utilicen esta celda para poder imprimir y explorar las anotaciones.
```

<!-- grade_id: cell-3b895917f8dcdbca, points: 0.2, grade: true, solution: false -->
```python
assert(type(dict_train)==dict), 'Las anotaciones deben ser un diccionario'
assert(type(dict_val)==dict), 'Las anotaciones deben ser un diccionario'
assert(type(dict_test)==dict), 'Las anotaciones deben ser un diccionario'

assert 'images' in dict_train.keys(), 'Las anotaciones deben tener la llave images'
assert 'images' in dict_val.keys(), 'Las anotaciones deben tener la llave images'
assert 'images' in dict_test.keys(), 'Las anotaciones deben tener la llave images'
```

<!-- grade_id: cell-b84394825072a5f1, points: 0.1, grade: true, solution: false -->
```python
assert(np.isclose(len(dict_train['images']), len(list_train))), 'La cantidad de imagenes de anotacion deben ser iguales que las de la base de datos'
assert(np.isclose(len(dict_val['images']), len(list_val))), 'La cantidad de imagenes de anotacion deben ser iguales que las de la base de datos'
assert(np.isclose(len(dict_test['images']), len(list_test))), 'La cantidad de imagenes de anotacion deben ser iguales que las de la base de datos'
```

## Parte 2: Visualización de imagenes

Como pudieron observar anteriormente y del contexto del problema, estamos tratando una tarea de detección. La intuición inicial sería basarnos en color, ya que los elementos se distinguen por la tinción que tienen en las imagenes. Para esto tenemos que, en primer lugar, poder observar las anotaciones exactamente donde se encuentran respecto a el objeto de interés y, en un segundo momento, hacer una exploración de los diferentes canales de color para así poder escoger el canal o los canales que mejor permitan distinguir el objeto de interes.

### 2.1 Visualización de anotaciones
Inicialmente es importante visualizar las anotaciones, esto se debe a que se está tratando con una tarea visual. Esto nos puede ayudar a ver que tan cercanas son las anotaciones a el objeto, es decir, en este caso, si está sobre el nucleo de la celula o se encuentra al rededor de la celula o incluso un poco mas de espacio al rededor esta. Para este proposito vamos a realizar un subplot de **4 filas por 2 columnas** donde la primera columna contenga las imagenes originales y las versiones anotadas de dichas imagenes. 
Recuerden usar ``fig, ax = plt.subplots(n_rows,n_cols)`` para hacerlo, manteneniendo ``ax``.

Primero, cree una función que reciba de input el diccionario de anotaciones y un path de la imagen a plotear. Debe retornar una imagen con la bounding box de la anotación correspondiente.

**Pasos a seguir:**

> - Lea la imagen del path de entrada y guardelo en una variable.

> - El diccionario de anotacion tiene una llave "images" que retorna unaa lista de diccionarios. Estos diccionarios tienen una llave "file_name" que retorna el nombre de la imagen y una llave "id" que retorna un numero entero. Encuentre el **image_id** que corresponda a la imagen del path de entrada.

> - El diccionario de anotaciones tiene una llave "annotations" que retorna una lista de diccionarios. Busque el diccionario cuya llave "image_id" sea igual al **id** de la imagen de entrada y con la llave "category_id" igual a 3 (la categoria de los globulos blancos).

> - De este ultimo diccionario, use la llave "bbox" para obtener una lista de 4 numeros [X,Y,W,H] donde X,Y son las coordenadas de la esquina superior izquierda de la bounding box y W,H son el tamaño del ancho y alto de la caja. 

> - Use la función cv2.rectangle(image, start_point, end_point, (255,0,0), 2) para obtener una imagen con una bounding box dibujada y retornela junto con las coordenadas del bbox. Recuerde que las coordenadas deben ser enteras.

<!-- grade_id: cell-35ea08541f527cd9, points: 0, grade: false, solution: true -->
```python
def Graficar_bbox(dict_anotaciones, path):
    # YOUR CODE HERE
    raise NotImplementedError()
    return img_bbox,bbox
```

<!-- grade_id: cell-a8b896e8fdc00572, points: 0, grade: false, solution: true -->
```python
plt.rcParams['figure.figsize'] = [10, 20] # por cuestiones de visualización vamos a alterar este parametro
fig, ax = plt.subplots(4,2)#filas,#columnas
ax[0,0].set_title('Imagen \n sin anotación')
ax[0,1].set_title('Imagen \n con anotación')

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cdc0320b5cb12295, points: 0.1, grade: true, solution: false -->
```python
for i in range(4):
    for j in range(2):
        assert(not ax[i,j].axison), 'Los ejes deben estar apagados'
assert(np.shape(ax) == (4,2)), 'EL subplot no es 4,2'
```

<!-- grade_id: cell-528ffacb57b785ea, points: 0.3, grade: true, solution: false -->
```python
bbox=Graficar_bbox(dict_train,list_train[0])[1]
assert bbox[0]==44, 'su coordenada de filas de la esquina es erronea'
assert bbox[1]==273, 'su coordenada de columnas de la esquina es erronea'
assert bbox[2]==142, 'El ancho de la caja es erroneo'
assert bbox[3]==142, 'El alto de la caja es erroneo'
```

### 2.2 Exploración por canales

Ahora que ya hemos visualizado las anotaciones tenemos conocimiento de como lucen estas y exactamente como se relacionan con el objeto, lo cual nos ayudará a realizar el proceso de forma automatica mas preciso. Sin embargo, como saben, la aproximación que vamos a realizar para poder resolver esta tarea es mediante segmentación y esta será basada en la información de color que tenemos disponible. Debido a esto, vamos a usar las 4 imagenes que usamos anteriormente y vamos a realizar un subplot con la siguiente forma:

<div>
<img src="canales.jpg" width="500"/>
</div>

> **consejos**: 
> - Recuerden que los subplots deben estar con sus respectivos titulos en cada fila para saber a que pertenece cada imagen. 
> - Tengan presente que deben utilizar el mapa de color de grises al presentar una imagen de un solo canal a no ser que se trate de una grafica. Para esto pueden usar el comando ``cmap`` al momento de graficar
> - Recuerden quitar los ejes de las imagenes

<!-- grade_id: cell-689313b5db94708c, points: 0, grade: false, solution: true -->
```python
plt.rcParams['figure.figsize'] = [10, 10] # por cuestiones de visualización vamos a alterar este parametro
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7c09d3359bc643d2, points: 0.1, grade: true, solution: false -->
```python
assert(np.shape(ax) == (4,4)), 'El subplot debe ser 4x4'

for i in range(4):
    for j in range(4):
        assert(not ax[i,j].axison), 'Los ejes deben estar apagados'
```

### 2.3 Umbrales

Como dijimos con anterioridad, usaremos un metodo de segmentación para resolver el problema de detección de globulos blancos. Hagamos una pequeña implementación de la segmentación de globulos basado en color utilizando las intensidades del canal azul.

Veamos ahora una manera más sencilla de umbralizar una imagen que utilizando un for. Para eso Utilice el operador > 215 para cambiar todos los pixeles mayores a 215 por un True y el resto por un False. Muestre la imagen resultante.

```python
# Por favor no modificar la imagen. Si la modifican pueden tener errores luego.
img=imread(list_train[0])
```

<!-- grade_id: cell-77c71c43116492e3, points: 0, grade: false, solution: true -->
```python
new_im='' #imagen umbralizada de la segunda forma

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-82f69aa388485ffd, points: 0.2, grade: true, solution: false -->
```python
assert np.isclose(np.sum(new_im),5708) or np.isclose(np.sum(new_im), 5597), 'La umbralización es erronea'
```

En el siguiente paso utilizaremos una función llamada [inRange](https://docs.opencv.org/3.4/d2/de8/group__core__array.html#ga48af0ab51e36436c5d04340e036ce981) de la librería OpenCV. Esta función nos permitira no solamente realizar una segmentación cómo la anterior si no combinarla en los diferentes canales de la imagen. La documentación la pueden encontrar en el siguiente [enlace](https://docs.opencv.org/3.4/d2/de8/group__core__array.html#ga48af0ab51e36436c5d04340e036ce981). Basicamente esta función utiliza dos arreglos o escalares que serán los límites inferiores y límites superiores para realizar la segmentación. Si la entrada son dos escalar (uno para el límite inferior y otro para el límite superior) entonces cada uno de los canales utiliza los los mismos valores. Mientras que si la entrada es dos arreglos cada canal tendrá un límite inferior y otro superior.

La idea es que realicen algunos experimentos y encuentren la mejor combinación de límites inferiores y superiores para la segmentación. Recuerden que hasta ahora no estamos utilizando métricas por lo que su respuesta debe basarse en los resultados cualitativos.

<!-- grade_id: cell-6b6c5bf5861f5bf5, points: 0, grade: false, solution: true -->
```python
segmentation = '' # Guarde acá su mejor experimento
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-204fb662acbf0441, points: 0.2, grade: true, solution: false -->
```python
assert np.sum(segmentation) > 0, 'La segmentación es erronea'
```

### Parte 3: Base de datos

Una vez se tienen disponibles los archivos con las imagenes y las anotaciones es importante hacer una revisión de la base de datos. Esto es importante ya que se debe conocer cuantas imagenes, cuantas anotaciones y cuales anotaciones  por imagen por carpeta y en total. Esta parte tiene una especial relevancia ya que en muchas ocasiones vamos a encontrarnos con bases de datos que tengan una distribución mal hecha (como pocos datos de entrenamiento) o estén desbalanceados, teniendo mas anotaciones en promedio en una carpeta. Tener pocas imagenes en entrenamiento o no tener una suficiente representación de una clase en alguna carpeta puede afectar el desempeño de nuestro modelo. 

Adicionalmente, debemos entender nuestro archivo de anotaciones. Como vimos con anterioridad, estos son diccionarios que almacenan la información necesaria para resolver nuestro problema de dedtección. Exploraremos la estructura de nuestro archivo de anotaciones y algunas estadisticas importantes.

<!-- grade_id: cell-c12f038e36876ca8, points: 0, grade: false, solution: true -->
```python
train_images = "" # variable que tendrá el numero de imagenes en la carpeta de entrenamiento
val_images = "" # variable que tendrá el numero de imagenes en la carpeta de validacion
test_images = "" # variable que tendrá el numero de imagenes en la carpeta de prueba

train_anots_1 = "" # variable que tendrá el numero de anotaciones de la clase 1 en la carpeta de entrenamiento
val_anots_1 = "" # variable que tendrá el numero de anotaciones de la clase 1 en la carpeta de validacion
test_anots_1 = "" # variable que tendrá el numero de anotaciones de la clase 1 en la carpeta de prueba

train_anots_2 = "" # variable que tendrá el numero de anotaciones de la clase 2 en la carpeta de entrenamiento
val_anots_2 = "" # variable que tendrá el numero de anotaciones de la clase 2 en la carpeta de validacion
test_anots_2 = "" # variable que tendrá el numero de anotaciones de la clase 2 en la carpeta de prueba

train_anots_3 = "" # variable que tendrá el numero de anotaciones de la clase 3 en la carpeta de entrenamiento
val_anots_3 = "" # variable que tendrá el numero de anotaciones de la clase 3 en la carpeta de validacion
test_anots_3 = "" # variable que tendrá el numero de anotaciones de la clase 3 en la carpeta de prueba

# YOUR CODE HERE
raise NotImplementedError()
```

Debemos comprobar si alguna imagen de nuestra base de datos no tiene la clase 3 ya que es nuestra clase de interes. A su vez, debemos comprobar cuantos objetos de nuestra clase de interes se encuentran en una imagen, si existe más de una celula blanca en nuestras imagenes, debemos ajustar nuestra metodología de detección.

El proceso de analsis de nuestras anotaciones puede ser algo complicado si nunca ha trabajado con diccionarios o un alto volumen de datos. Les proponemos el siguiente pseudo-codigo.

**Pasos:** 

> - La llave "anotaciones" del json retorna una lista de dissionarios. Estos diccionarios tienen un "image_id" que es el identificador de la imagen y un "category_id" que es la clase del objeto. Haga una lista que contenga solo los diccionarios de la category id =3. 

> - Haga una lista de 0s con np.zeros() con el tamaño de la cantidad de imagenes de entrenamiento (o validacion o test).

> - Recorra todos los diccionarios del primer paso y guarde el image_id en una variable. Este image_id es un entero.

> - Sume 1 en la posición del image_id en la lista de 0s. De esta forma sabremos cuantos objetos de la categoria 3 se encuentran en cada imagen sabiendo que la image_id=5 tendrá la posición 5 en nuestro vector de 0s (y así con todas las imagenes).

> - Use np.min() o np.max() en el vector de 0s.

<!-- grade_id: cell-e58f515eedcdede0, points: 0, grade: false, solution: true -->
```python
min_train_anots_3= "" # variable que tendrá el minimo de elementos de la clase 3 en una imagen en entrenamiento
min_val_anots_3= "" # variable que tendrá el minimo de elementos de la clase 3 en una imagen en validacion
min_test_anots_3= "" # variable que tendrá el minimo de elementos de la clase 3 en una imagen en prueba

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1b9d7836415ed1f, points: 0, grade: false, solution: true -->
```python
max_train_anots_3= "" # variable que tendrá el maximo de elementos de la clase 3 en una imagen en entrenamiento
max_val_anots_3= "" # variable que tendrá el maximo de elementos de la clase 3 en una imagen en validacion
max_test_anots_3= "" # variable que tendrá el maximo de elementos de la clase 3 en una imagen en prueba

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b1bbf224c4ba0cc, points: 0.2, grade: true, solution: false -->
```python
assert type(train_images)==int, 'El numero de imagenes de entrenamiento deberia ser entero'
assert type(val_images)==int, 'El numero de imagenes de validacion deberia ser entero'
assert type(test_images)==int, 'El numero de imagenes de test deberia ser entero'

assert type(train_anots_1)==int, 'El numero de anotaciones de entrenamiento deberia ser entero'
assert type(val_anots_1)==int, 'El numero de anotaciones de validacion deberia ser entero'
assert type(test_anots_1)==int, 'El numero de anotaciones de test deberia ser entero'

print ("{:<8}             | {:<15} {:<20}".format('Fold','Total images','Total annotations'))
print ('----------------------------------------------------------------------------')
print ("{:<20} | {:<15} {:<20}".format('Train category 1',train_images,train_anots_1))
print ("{:<20} | {:<15} {:<20}".format('Valid category 1',val_images,val_anots_1))
print ("{:<20} | {:<15} {:<20}".format('Test category 1',test_images,test_anots_1))

assert train_images<train_anots_1, 'Deberían haber mas anotaciones de la clase 1 en entrenamiento que el total de imagenes'
assert val_images<val_anots_1, 'Deberían haber mas anotaciones de la clase 1 en entrenamiento que el total de imagenes'
assert test_images>test_anots_1, 'Deberían haber menos anotaciones de la clase 1 en entrenamiento que el total de imagenes'

assert np.isclose(train_anots_1*val_anots_1*test_anots_1,397800), 'Algun valor de la cantidad de anotaciones esta mal'
```

<!-- grade_id: cell-d8a34ba578e91711, points: 0.2, grade: true, solution: false -->
```python
assert type(train_anots_2)==int, 'El numero de anotaciones de entrenamiento deberia ser entero'
assert type(val_anots_2)==int, 'El numero de anotaciones de validacion deberia ser entero'
assert type(test_anots_2)==int, 'El numero de anotaciones de test deberia ser entero'

print ("{:<8}             | {:<15} {:<20}".format('Fold','Total images','Total annotations'))
print ('----------------------------------------------------------------------------')
print ("{:<20} | {:<15} {:<20}".format('Train category 2',train_images,train_anots_2))
print ("{:<20} | {:<15} {:<20}".format('Valid category 2',val_images,val_anots_2))
print ("{:<20} | {:<15} {:<20}".format('Test category 2',test_images,test_anots_2))

assert train_images<train_anots_2, 'Deberían haber mas anotaciones de la clase 2 en entrenamiento que el total de imagenes'
assert val_images<val_anots_2, 'Deberían haber mas anotaciones de la clase 2 en entrenamiento que el total de imagenes'
assert test_images<test_anots_2, 'Deberían haber mas anotaciones de la clase 2 en entrenamiento que el total de imagenes'

assert np.isclose(train_anots_2*val_anots_2*test_anots_2,670077408), 'Algun valor de la cantidad de anotaciones esta mal'
```

<!-- grade_id: cell-05c752532d063975, points: 0.2, grade: true, solution: false -->
```python
assert type(train_anots_3)==int, 'El numero de anotaciones de entrenamiento deberia ser entero'
assert type(val_anots_3)==int, 'El numero de anotaciones de validacion deberia ser entero'
assert type(test_anots_3)==int, 'El numero de anotaciones de test deberia ser entero'

print ("{:<8}             | {:<15} {:<20} {:<20} {:<20}".format('Fold','Total images','Total annotations', 'Min annotations', 'Max annotations'))
print ('----------------------------------------------------------------------------')
print ("{:<20} | {:<15} {:<20} {:<20} {:<20}".format('Train category 3',train_images,train_anots_3,min_train_anots_3,max_train_anots_3))
print ("{:<20} | {:<15} {:<20} {:<20} {:<20}".format('Valid category 3',val_images,val_anots_3,min_val_anots_3,max_val_anots_3))
print ("{:<20} | {:<15} {:<20} {:<20} {:<20}".format('Test category 3',test_images,test_anots_3,min_test_anots_3,max_test_anots_3))

assert train_images<train_anots_3, 'Deberían haber mas anotaciones de la clase 3 en entrenamiento que el total de imagenes'
assert val_images==val_anots_3, 'Deberían haber igual cantidad de anotaciones de la clase 3 en entrenamiento que el total de imagenes'
assert test_images<test_anots_3, 'Deberían haber mas anotaciones de la clase 3 en entrenamiento que el total de imagenes'

assert np.isclose(train_anots_3*val_anots_3*test_anots_3,486720), 'Algun valor de la cantidad de anotaciones esta mal'

assert min_train_anots_3+min_val_anots_3+min_test_anots_3==2, 'Algun valor de minimo esta mal'
assert max_train_anots_3+max_val_anots_3+max_test_anots_3==6, 'Algun valor de maximo esta mal'
```

## Entregables
Los entregables para esta primera entrega son los siguientes:

- Jupyter notebook (.ipynb): El notebook debe estar completamente resuelto. Recuerden que la nota depende del correcto funcionamiento de su código.

- Archivo de texto (.txt): Un archivo de texto con las secciones de código del Jupyter notebook. Para más información sobre cómo convertir las secciones de código de su Jupyter notebook a un archivo de texto, pueden ver el video disponible en el siguiente [enlace](https://uniandes-edu-co.zoom.us/rec/share/QOxUUIw7Uz9DsnKPyftOXZgM4bx5d7KHMYK-dqvGnkyJ2GfqvT7NU3lhQc0NtNSU.jveWB2H34S4uLIoB).

- Informe (.pdf): El informe constará de responder las preguntas, además, asegúrese de incluir los elementos adicionales, como subplots, gráficos, tablas o cualquiera elemento adicional que se mencione. Toda la información sobre preguntas y elementos adicionales está en la siguiente sección de informe. El informe debe utilizar el formato que puede encontrar siguiente [enlace](https://github.com/cvpr-org/author-kit). Si tiene dudas de cómo utilizar el formato puede preguntar al Asistente Graduado.

## Informe 
La extensión máxima es de 4 páginas. Las preguntas a responder en su informe son las siguientes.

### Revisión de estadisticas
Una parte importante en toda tarea de este campo, es no solo observar los problemas que puedan presentarse en la base de datos, sino poder comprender completamente cual es el proposito de cada uno de los elementos y por que deben cumplir con ciertas caracteristicas. Para poder comprender esto usted debe estar en la capacidad de responder las siguientes preguntas, por lo que debe responderlas en su informe:

- ¿Cuál es la función de cada uno de los sub-conjuntos de una base de datos? 
- ¿Hay una clase en la base de datos que es predominante sobre las otras? ¿Si la hay cuál es esa clase?
- ¿La base de datos está distribuida correctamente en cada uno de los set de datos (train, validación y test)? **Consejo:** Revisar si el porcentaje de cada categoría es similar en cada set de datos.
- En la base de datos hay algunas imágenes que no tienen ningún glóbulo blanco. Dado que lo que queremos detectar son estos últimos ¿estas imágenes, sin glóbulos blancos, tienen alguna función dentro de la base de datos o su utilidad es nula?
- Realice la presentación de las estadísticas de la base de datos de una oportuna en el informe. Recuerde que puede utilizar diversos métodos de visualización como los diagramas de barras, histogramas, entre otros.
- Finalmente deben realizar un gráfico que muestre la organización de las anotaciones (archivo _annotations.coco.json). Es decir, deben enunciar que elementos (keys) tienen y que contiene cada elemento. Además en caso de que los elementos al interior también sean diccionarios, deben enunciar los elementos de estos. Explique en detalle que es el image_id y el category_id ¿Cuál es su utilidad?.

### Visualización de imágenes
Como deben recordar, uno de los propositos de esta entrega era aproximarnos tambien al problema en si. Debido a esto deben responder las siguientes preguntas sobre este problema y las imágenes a usar.

- ¿Qué canal parece tener mejores resultados? ¿Utilizar varios canales para hacer una umbralización final puede dar mejores resultados? Explique su razonamiento.
- Presente los resultados de los experimentos más importantes que realizó con la función inRange. Realice un análisis de esos resultados.
- Investigue un algoritmo sencillo que pueda ser utilizado para resolver esta tarea. Deben ser técnicas que utilicen lo aprendido hasta ahora o todo lo utilizado en el primer módulo del curso. **Nota**: Recuerden que hasta ahora no hemos resuelto la tarea final que es detección. Solo hemos hecho segmentación con la intención de que esto nos ayude a detectar los glóbulos blancos.
- Finalmente deben poner cada uno de los subplots que realizamos, acompañados de una breve descripción de estos. Recuerden que una descripción es diferente a un análisis.

```python
from utils import *
converter("Entrega 1")
```